# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return (result)
    #raise NotImplementedError()
example_word_count()
```

Output:
```
['Amy', 'is', '5', 'years', 'old']
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""
    names = re.findall(r'([A-Z][a-z]+)',simple_string)
    return(names)
    # YOUR CODE HERE
    raise NotImplementedError()
names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [grades.txt](grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("grades.txt", "r") as file:
        grades = file.read()
    grade = re.findall(r'([A-Z]\w+\s[A-Z]\w+:\sB)',grades)
    # YOUR CODE HERE
    return(grade)
    raise NotImplementedError()
grades()
```

Output:
```
['Bell Kassulke: B',
 'Simon Loidl: B',
 'Elias Jovanovic: B',
 'Hakim Botros: B',
 'Emilie Lorentsen: B',
 'Jake Wood: B',
 'Fatemeh Akhtar: B',
 'Kim Weston: B',
 'Yasmin Dar: B',
 'Viswamitra Upandhye: B',
 'Killian Kaufman: B',
 'Elwood Page: B',
 'Elodie Booker: B',
 'Adnan Chen: B',
 'Hank Spinka: B',
 'Hannah Bayer: B']
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16
```

## Part C

Consider the standard web log file in [logdata.txt](logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("logdata.txt", "r") as file:
        logdata = file.read()
    # lines = re.split("\n+",logdata)
    # ip = [re.split("(\d+\.)*\d+\s-\s",ip,len(lines))for ip in lines]
    # host = re.findall("\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3}",logdata)
    # user_name = re.findall("([a-z]+\d+)",logdata)
    # time = re.findall("\d{1,2}/Jun/2019:\d{1,2}:\d{1,2}:\d{1,2} -0700",logdata)
    # request = re.findall("[A-Z]+\s/.*\sHTTP/\d\.\d",logdata)
    # return(request)
    # pattern = """
    # (?P<host>\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3})
    # (\s-\s)
    # (?P<username>([a-z]+\d+)|(-))
    # (\s\[)
    # (?P<time>\d{1,2}/Jun/2019:\d{1,2}:\d{1,2}:\d{1,2} -0700)
    # (\])
    # (\s\")
    # (?P<request>[A-Z]+\s/.*\sHTTP/\d\.\d)
    # (\.*)
    # """
    for item in re.finditer("(?P<host>\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3})(\s-\s)(?P<username>([a-z]+\d+)|(-))(\s\[)(?P<time>\d{1,2}/Jun/2019:\d{1,2}:\d{1,2}:\d{1,2} -0700)(\])(\s\")(?P<request>[A-Z]+\s/.*\sHTTP/\d\.\d)(\.*)",logdata):
        return item.groupdict()
logs()
```

Output:
```
{'host': '30.95.91.251',
 'username': 'larson8319',
 'time': '21/Jun/2019:16:02:02 -0700',
 'request': 'PUT /one-to-one/whiteboard HTTP/1.0'}
```

```python
# with open("logdata.txt", "r") as file:
#         logdata = file.read()
```

```python
# pattern = """
#     (?P<host>\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3})
#     (\s-\s)
#     (?P<username>([a-z]+\d+)|(-))
#     (\s\[)
#     (?P<time>\d{1,2}/Jun/2019:\d{1,2}:\d{1,2}:\d{1,2} -0700)
#     (\])
#     (\s\")
#     (?P<request>[A-Z]+\s/.*\sHTTP/\d\.\d)
#     (\.*)
#     
#     """
# for item in re.finditer(pattern,logdata):
#         print(item.groupdict()['host'])
# for host in re.findall("\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3}",logdata):
#     print(host)
# for item in re.finditer("(?P<host>\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3})(\s-\s)(?P<username>([a-z]+\d+)|(-))(\s\[)(?P<time>\d{1,2}/Jun/2019:\d{1,2}:\d{1,2}:\d{1,2} -0700)(\])(\s\")(?P<request>[A-Z]+\s/.*\sHTTP/\d\.\d)(\.*)",logdata):
#     print(item.groupdict())
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979
one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In [7], line 1
----> 1 assert len(logs()) == 979
      2 one_item={'host': '146.204.224.152',
      3   'user_name': 'feest6811',
      4   'time': '21/Jun/2019:15:45:24 -0700',
      5   'request': 'POST /incentivize HTTP/1.1'}
      6 assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

TypeError: object of type 'NoneType' has no len()
```